## Task 1: Encoding and Decoding Text

We will write simple functions to encode and decode text.  

### The decoding function

The decoding function takes as input a list of $n$ words $[w_0, w_1, \ldots, w_{n-1}]$, and a list of $m$ integers $[k_0, k_1, \ldots, k_{m-1}]$, where each integer is in the range $[0, \ldots, n-1]$.

The function must output the string $w_{k_0} \; w_{k_1} \; w_{k_2} \; \ldots w_{k_m}$, where each number has been replaced with the word in the corresponding position in the list, and words are joined by spaces.

For example, assume the function is called with:

    textdecode(["a", "bird", "seed", "eats"], [0, 1, 3, 0, 2])
    
Then the output should be

    "a bird eats a seed"
    
because word 0 in the list is "a", word 1 is "bird", word 3 is "eats", word 0 is "a", and word 2 is "seed".

Here, we remind you that if you have a list of word, you can join them using spaces to form a string like this:

    " ".join(["coffee", "is", "good"])
    
which produces

    "coffee is good"


```python
def textdecode(words, indices):
    """Decodes, producing a sentence in which each index is replaced
    by the word in the corresponding position in words, and the words
    are joined via spaces."""
    sentance = []
    for i in indices:
      sentance.append(words[i])

    k = " ".join(sentance)

    return k
```

```python
# Cell for testing code
'''
def textdecode(words, indices):
  sentance = []
  for i in indices:
    sentance.append(words[i])
  for x in sentance:
    " ".join(sentance)
  for y in sentance:
    print(y, end=' ')

'''
```

```python
# Test cases

assert textdecode(["a", "bird", "seed", "eats"], [0, 1, 3, 0, 2]) == "a bird eats a seed"
assert textdecode(["ha"], [0, 0, 0]) == "ha ha ha"
assert textdecode(["be", "to", "not", "or"], [1, 0, 3, 2, 1, 0]) == "to be or not to be"
```

### Encoding function

The encoding function works in the opposite way.  Given a sentence `s`, it first splits `s` according to spaces, via `s.split()`, obtaining a list of words.  It then removes duplicates from this list of words, and returns:

* The list of words, with duplicates removed;
* A list obtained by replacing each word with the index of the word in the above list.

For instance, if the input is:

    "i need a coffee"
    
then the words are `"i"`, `"need"`, `"a"`, `"coffee"`, and there are no duplicates.  The output will then be:

    ["i", "need", "a", "coffee"], [0, 1, 2, 3]
    
If the [input is](https://www.youtube.com/watch?v=98luuj8ymoQ)

    "yo ho ho and a bottle of rum"
    
the output can be

    ["yo", "ho", "and", "a", "bottle", "of", "rum"], [0, 1, 1, 2, 3, 4, 5, 6]
    
or if you prefer,

    ["a", "bottle", "of", "rum", "and", "yo", "ho"], [5, 6, 6, 4, 1, 2, 3]
    

<!-- grade_id: cell-c475a252851c0241, points: 0, grade: false, solution: true -->
```python
def textencode(s):
    """Encode the string s, dividing it into words, and returning a list
    of unique words, and the list of indices of each word of s in the word list,
    in order."""
    list_of_words = s.split( )



    no_repeat = []
    for i in list_of_words:
      if i not in no_repeat:
        no_repeat.append(i)



    index_of_sentance = []

    for x in list_of_words:
      k = no_repeat.index(x)

      index_of_sentance.append(k)



    return no_repeat, index_of_sentance


```

```python
# some simple tests

assert textencode("a") == (["a"], [0])
assert textencode("ho ho ho") == (["ho"], [0, 0, 0])

```

<!-- grade_id: cell-c5512e5a9a318fbd, points: 10, grade: true, solution: false -->
```python
# Encoding / decoding tests

import random

words = ["a", "bi", "ci", "di", "e", "effe", "gi"]
for _ in range(1000):
    k = random.randint(0, 10)
    s = " ".join(random.choices(words, k=k))
    ws, ix = textencode(s)
    assert len(set(ws)) == len(set(ix))
    t = textdecode(ws, ix)
    assert s == t


```

## Task 2: Numerical Dictionaries

This task considers _numerical dictionaries_, which map keys to integers (rather than to arbitrary values).  These dictionaries can be used to count occurrences; for instance, the dictionary {'a': 1, 'b': 3} indicates that 'a' occurred once, and 'b' three times.

Write a function `add_dicts` that, given two numerical dictionaries, computes the numerical dictionary resulting from their sum.  For instance, `add_dict({'a': 1, 'b': 3}, {'b': 2, 'c': 3})` should give `{'a': 1, 'b': 5, 'c': 3}` as result.

<!-- grade_id: cell-ea9265bbf770b82a, points: 0, grade: false, solution: true -->
```python
def add_dicts(d1, d2):
  d3 = {}
  for c,d in d1.items():
    d3[c] = d


  for a, b in d3.items():
    for x, y in d2.items():
      if a == x:
        d3[a] = (b+y)
      #else:
        #d3[x] = y
  return d3
```

<!-- grade_id: cell-3cafd3159a927abc, points: 10, grade: true, solution: false -->
```python
# Test for when the keys are the same.

assert add_dicts({'a': 1, 'b': 2}, {'a': 3, 'b': 1}) == {'a': 4, 'b': 3}
# And empty.
assert add_dicts({}, {}) == {}


```

<!-- grade_id: cell-1fc7a9ffa701cef9, points: 10, grade: true, solution: false -->
```python
# Test for when the keys are different.

assert add_dicts({'a': 1, 'c': 2}, {'a': 3, 'b': 1}), {'a': 4, 'b': 1, 'c': 2}
assert add_dicts({'a': 3, 'b': 1}, {'a': 1, 'c': 2}), {'a': 4, 'b': 1, 'c': 2}


```

## Task 3

Implement a function `ordered_substring` that takes as input two strings, s and t.
The function returns True if and only if the characters of s can be found, in the order in which they are in s, as characters of t. In other words, the function returns True if and only if it is possible to remove some characters from t, so that the remaining ones  form the string s.

For example:

* `ordered_substring("abc", "a123b4b523bc")` should return True, because `"a123b4b523bc"` contains the characters `"abc"` in that order, taking its 1st, 5th, and last character.
* Instead, `ordered_substring("abc", "a123c4b523b")` should return False, because `"a123c4b523b"` does not contain a "c" after a "b".
* `ordered_substring("abb", "a123c4b523f")` should also return False, because there is no `"b"` after the first `"b"` in `"a123c4b523f"`.


<!-- grade_id: cell-de576dac776780ee, points: 0, grade: false, solution: true -->
```python
def ordered_substring(s, t):
    """Returns True if s is an ordered substring of t, and False otherwise."""
    letters_check = []
    for i in s:
      letters_check.append(i)
    '''
    letters_total = []
    for a in t:
      letters_total.append(a)
  '''

    j = 0
    for x in letters_check:
      letter_index = t.find(x,j)
      if letter_index == -1:
        return False
      j = letter_index + 1
    return True



ordered_substring('auu', 'cabd')
```

Output:
```
False
```

<!-- grade_id: cell-7730e72d0b13dc13, points: 10, grade: true, solution: false -->
```python
# Test that the code works for empty and one-character strings s.

assert ordered_substring('', 'abc')
assert ordered_substring('a', 'cabd')
assert ordered_substring('d', 'cabd')
assert not ordered_substring('f', 'cabd')
assert not ordered_substring('a', '')
assert ordered_substring('', '')


```

<!-- grade_id: cell-b6dc190dc3e32776, points: 10, grade: true, solution: false -->
```python
# Some simple cases.

assert ordered_substring('ab', 'abcd')
assert ordered_substring('ac', 'abcd')
assert ordered_substring('bd', 'abcd')
assert not ordered_substring('ba', 'abcd')
assert not ordered_substring('af', 'abcd')
assert ordered_substring('123', '12345')


```

<!-- grade_id: cell-7757e7c597c838f2, points: 10, grade: true, solution: false -->
```python
# General cases.

assert ordered_substring('abc', 'abbbeec')
assert not ordered_substring('abbc', 'abc')
assert ordered_substring('aa', 'bacad')
assert ordered_substring('aaba', 'ababa')


```

## Task 4: Elections

Write an `Election` class, used to record votes.  The class has the following methods:

* `vote(c)`: records a vote for candidate `c`.
* `total_votes()`: returns the total number of votes.
* `winner()`: returns the set of candidates with maximal votes.  Normally there is only one, but there could be a tie, in which case you return the set of tied candidates.  Return `None` if no votes have been cast.
* `margin()`: returns the margin of victory, that is, the number of votes separating the fist from the second candidate.  This can be 0, in case of a tie.  If there is only one candidate, you can answer with the number of votes for the candidate.  If no votes have been cast, return `None`.



<!-- grade_id: cell-b4a77cb854779176, points: 0, grade: false, solution: true -->
```python
class Election(object):
    """Implements the election class."""

    def __init__(self):
        # Dictionary mapping voters to the number of votes received.
        self.votes = {}

    def vote(self, c):
        """Records a vote for candidate c."""
        new_vote = self.votes.get(c,0)
        self.votes[c] = new_vote + 1

    def total_votes(self):
        """Returns the total number of votes cast."""
        total = 0
        for v in self.votes.values():
          total += v
        return total

    def winner(self):
        """Returns the list of winners of the election."""
        max = 0
        total_winners = []
        if len(self.votes) == 0:
          return None
        else:
           for i, j in self.votes.items():
            if j >= max:
              max = j
              winner_vote = i
              total_winners.append(winner_vote)
        print("total_winners: ", total_winners)


        return total_winners


    def margin(self):
        """Returns the margin of victory of the winner."""
        runner_up = 0
        top = 0
        if len(self.votes) == 0:
          return None
        else:
          for p in self.votes.values():
            print(p)
            if p > top:
              runner_up = top
              top = p
            elif p > runner_up:
              runner_up = p


        return top - runner_up
```

<!-- grade_id: cell-51d556dc9c8d51e4, points: 10, grade: true, solution: false -->
```python
# Tests for margin method.

e = Election()
assert e.margin() == None

e.vote('a')
e.vote('b')
e.vote('a')
assert e.margin() == 1

e.vote('c')
e.vote('b')
assert e.margin() == 0


```

Output:
```
2
1
2
2
1
```

<!-- grade_id: cell-f41fa91bb1842bbb, points: 10, grade: true, solution: false -->
```python
# Tests for total_votes method.

e = Election()
assert e.total_votes() == 0

e.vote('a')
e.vote('b')
e.vote('a')
assert e.total_votes() == 3

e.vote('c')
e.vote('b')
assert e.total_votes(), 5


```

<!-- grade_id: cell-bfe4cd563a6bd6e6, points: 10, grade: true, solution: false -->
```python
# Tests for winner method.

e = Election()
assert e.winner() == None

e.vote('a')
e.vote('b')
e.vote('a')
assert set(e.winner()) == {'a'}

e.vote('c')
e.vote('b')
assert set(e.winner()) == {'a', 'b'}


```

Output:
```
total_winners:  ['a']
total_winners:  ['a', 'b']
```